<a href="https://colab.research.google.com/github/Ulugbek9403/ml_edu/blob/master/Example_test_(9).ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Привет!

Перед началом прохождения убедитесь, что все настроено как надо. Следующая ячейка должна выполниться без ошибок:

После прохождения не забудьте **перезапустить ядро** (**restart the kernel**) (в панели меню выберите Kernel $\rightarrow$ Restart) и затем **запуск всех ячеек** (**run all cells**) (в панели меню выберите Cell $\rightarrow$ Run All).

Убедитесь, что заменили все такие места
```python
# YOUR CODE HERE
raise NotImplementedError()
```
и
```
"YOUR ANSWER HERE"
```
на свои решения и ответы.

Если не установлены какие-то пакеты, то посмотрите нужные версии в ноутбуках по регрессии.

Для прорешивания используйте среды с поддержкой формата Jupyter: [Google Colab](https://colab.research.google.com/), [Jupyter Notebook (Web)](https://jupyter.org/), [PyCharm](https://www.jetbrains.com/ru-ru/pycharm/), [DataSpell](https://www.jetbrains.com/ru-ru/dataspell/) или [VSCode](https://code.visualstudio.com/). Работа с этим форматом в обычном текстовом редакторе может сыграть злую шутку.

Не забывайте о следующих особенностях:
- 😉 Помимо указанных тестов после задания существуют и скрытые тесты, поэтому внимательнее читайте задание
- 😎 Обязательно проверяйте выполнимость на тестах, для проверки перезапустите ядро и запустите все ячейки заново!
- 😂 После скачивания ноутбуков достаточно заполнить места для ответов и отправить ноутбук обратно (сам ноутбук, ни в архиве, ни ссылкой на GitHub ни как-то еще принимать не будем, загружаете сам файл)
- 😈 Не рекомендуем удалять исходные ячейки, так как из-за этого могут не начислиться баллы =(
  - 👼 Если таки такое случилось, скачайте заново шаблон и перенесите в него свои решения
- 😛 Свой код пишите только в ячейках для решений - любые дописывания в остальных ячейках (с текстом или тестами) могут не попасть в проверку (например, `import` вне ячейки для решения)


---

# Приветствие

Привет, в этом тесте вы познакомитесь с тремя примерами задач, которые будут у вас в будущих контрольных точках.

Задачи не сложные и нацелены лишь на то, чтобы просто вы опробовали свои силы и познакомились с форматом!

## Задание 1 (1 балл)

Напишите код разворачивания списка.

Вместо блока
```python
# YOUR CODE HERE
raise NotImplementedError()
```
нужно реализовать свой код, чтобы он проходил по следующим тестам.

Сразу учитывайте, что не все тесты открыты и часть из них будет проверяться после сдачи решения.

> Решение пиши сразу в ячейке ниже

<!-- grade_id: cell-b7df227e69ae6723, points: 0, grade: false, solution: true -->
```python
def reverse_list(data: list):
    """
    Разворачивает список в обратном порядке.

    Parameters:
        data (list): Исходный список.

    Returns:
        list: Развернутый список.
    """
    return data[::-1]
    raise NotImplementedError()
```

А это открытые тесты, которые будут проверять решение. Не бойтесь дополнять их, если хотите проверить больше кейсов.

<!-- grade_id: cell-a9075953f720b87b, points: 1, grade: true, solution: false -->
```python
assert reverse_list([0]) == [0]
assert reverse_list([3, 2, 1]) == [1, 2, 3]
assert reverse_list([]) == []

```

## Задание 2 (1 балл)

Разработайте функцию вычисления площади круга по переданному радиусу. В случае переданного отрицательного радиуса должен вернуться 0.

> Здесь тоже решение пиши сразу в ячейке ниже

<!-- grade_id: cell-e0913548aff9db16, points: 0, grade: false, solution: true -->
```python
import math as m

def compute_circle_area(radius: float):
    """
    Вычисляет площадь круга по переданному радиусу.
    Возвращает 0, если радиус отрицательный.

    Parameters:
        radius (float): Радиус круга.

    Returns:
        float: Площадь круга или 0, если радиус отрицательный.
    """
    if radius < 0:
        return 0
    return m.pi * radius**2
    raise NotImplementedError()
```

<!-- grade_id: cell-4302ce285797495b, points: 1, grade: true, solution: false -->
```python
assert m.isclose(compute_circle_area(3), 28.274, abs_tol=1e-3)
assert m.isclose(compute_circle_area(-2), 0)

```

## Задание 3 (2 балла)

Ответьте кратко на вопрос.
Какие бывают типы данных Immutable, а какие Mutable? В чем разница между Immutable и Mutable типами?

> Надеюсь, и так понятно? =)  
> Решение пиши сразу в ячейке ниже

<!-- grade_id: cell-5008a206312d2427, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Задание 4 (0 баллов)

Напишите хокку

## Подзравляем!

Молодец, тебе удалось прорешать пример контрольной точки!

Теперь прогони ноутбук еще разок с перезапуском ядра (Run All - Такие стрелочки на панели сверху ">>" или меню "Kernel -> Restart & Run All") и убедись, что код работает верно и все открытые тесты проходят!